---
title       : User defined functions
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:01de87a8ac
## Simple functions


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
